# Homework 1 Assignment


The places where you have enter code are marked with `# YOUR CODE HERE`.

## Question 1 (8 points)

Complete the function `OR3` that implements the 3-qubit to 1-qubit $OR$ operator. This operator acts on 4 qubits in the state $|q_0,q_1,q_2,q_3\rangle$ and maps it to $|q_0,q_1,q_2,q_3 \oplus (q_0 \text{ or } q_1 \text{ or } q_2)  \rangle$.

The function `OR3` has

- Inputs: None
- Returns: A `QuantumCircuit` object with 4 qubits and 1 classical bit that implements the $OR$ operator as described above.

Don't add any measurements to the returned circuit.

<!-- grade_id: cell-517fe01d3e0bd3be, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumRegister, ClassicalRegister, QuantumCircuit
from qiskit.circuit.library.standard_gates import XGate
from qiskit_aer import AerSimulator


def OR3():
    circuit = QuantumCircuit(4,1)
    # Do not modify anything above this line
    
    # Dj was here
    circuit.x([0, 1, 2])
    
    circuit.mcx([0, 1, 2],3)

    circuit.x([0, 1, 2])
    
    circuit.x(3)
    # Do not modify anything below this line  
    return circuit
```

### Dj notas para no perdeme

> yo que resuelve esto recuerda que debes aplicar lo siguiente segun lo que dice la pregunta 1
> cosas a realizar:

 - en la entrada debo de recibir un 000 el cual devo invertir y convertir en 111
 - elmcx debe de activarse y cambiarme el q3 a 1 como lo hago?
 - invertor el q3 de 1 a 0
 - aplicar la compuerta OR

In the following cell, there is some code that will help you make sure your solution is correct. You can modify this code as you wish. You may also create more cells and write more code to test your own solution. 

```python
from qiskit.quantum_info import Statevector
import matplotlib.pyplot as plt
import cirq
import numpy as np

def test(input):
    circuit = QuantumCircuit(4,1)

    # implement the input state with a suitable use of X gates.
    for i in range(3):
        if input[i]=='1':
            circuit.x(i)

    # Apply the OR3 gate you created
    circuit.compose(OR3(), range(4), [0], inplace=True)

    circuit.measure(3,0)

    job = AerSimulator().run(circuit,shots=100)
    return job.result().get_counts(circuit)
test('001')

#ver si el cambio es correcto
def grafica_de_desarrollo(input_str):
    circuit = QuantumCircuit(4,1)
    for i in range(3):
        if input_str[i] == '1':
            circuit.x(i)
    estado_antes = Statevector.from_instruction(circuit)
    circuit.compose(OR3(), range(4), [0], inplace=True)
    estado_despues = Statevector.from_instruction(circuit)

    print(f"--- Estado antes de OR3 para input {input_str}: ---")
    display(estado_antes)
    print(f"Estado despues de OR3:")
    display(estado_despues)
    print("\n" + "="*50 + "\n")

grafica_de_desarrollo('001')


def estructura_circuito():
    circuit = cirq.Circuit()
    qq = cirq.LineQubit.range(3)

    circuit.append(cirq.H.on_each(*qq))

    # 3. El Oráculo: Reemplazamos 'a' con 0
    # Esto crea una compuerta CZ donde qq[0] es el control y qq[2] es el objetivo.
    CCCZ = cirq.Z(qq[2]).controlled_by(qq[0])
    circuit.append(CCCZ)

    # --- Impresión del circuito para análisis visual ---
    print("--- Estructura del Circuito Cuántico ---")
    print(circuit)
    print("\n" + "="*50 + "\n")

estructura_circuito()

test('001')

```

Output:
```
File ~/Documentos/Programacion/Qubit_MX/env/lib/python3.12/site-packages/cirq/ops/common_gates.py:1186, in CZPowGate.__repr__(self)
   1184     if self._exponent == 1:
   1185         return 'cirq.CZ'
-> 1186     return f'(cirq.CZ**{proper_repr(self._exponent)})'
   1187 return (
   1188     'cirq.CZPowGate('
   1189     f'exponent={proper_repr(self._exponent)}, '
   1190     f'global_shift={self._global_shift!r})'
   1191 )

File ~/Documentos/Programacion/Qubit_MX/env/lib/python3.12/site-packages/cirq/_compat.py:177, in proper_repr(value)
    174 if isinstance(value, np.ndarray):
    175     return f'np.array({value.tolist()!r}, dtype=np.{value.dtype!r})'
--> 177 if isinstance(value, pd.MultiIndex):
    178     return f'pd.MultiIndex.from_tuples({repr(list(value))}, names={repr(list(value.names))})'
    180 if isinstance(value, pd.Index):

AttributeError: module 'pandas' has no attribute 'MultiIndex'
```

```python

```

# How autograding will work in this workshop
In the next cell, there are some tests (the `assert` statements) that will be used to *grade* if your solution is correct. If executing the next cell results in an error, your implementation is incorrect, and you will not get any marks. 

There are also some hidden tests, not available to you, that will test your solution, once your submit this file back on Canvas. If any of those fail, you will not receive any marks. Otherwise, you will gain full marks on this question.

<!-- grade_id: cell-12b1e3e97b88002b, points: 2, grade: true, solution: false -->
```python
def Grade_test(input):
    circuit = QuantumCircuit(4,1)

    for i in range(3):
        if input[i]=='1':
            circuit.x(i)

    circuit.compose(OR3(), range(4), [0], inplace=True)
    circuit.measure(3,0)

    job = AerSimulator().run(circuit,shots=100)
    return job.result().get_counts(circuit)

counts = Grade_test('100')
assert counts['1']==100
counts2 = Grade_test('000')
assert counts2['0']==100
```

<!-- grade_id: cell-74f92b92e2a40b2e, points: 6, grade: true, solution: false -->
```python
# hidden tests in this cell will be used for grading.
```